```python
# Initialize Otter
import otter
grader = otter.Notebook("classesMethods.ipynb")
```

## Lecture Section

We have all the tools we need to talk about classes and objects. Both are an important piece of Object-Oriented-Programming (OOP) - a coding paradigm that structures software into reusable, modularized pieces of code. We won't cover OOP here, because Python does not allow for the full OOP paradigm. However, regardless of if you choose to create or make classes, you will (and already have) use objects very regularly. Understanding how they work will help you fix errors, understand behavior, and feel more confident in exploring and using libraries other than what we cover in this course.

In this lecture, we will cover:
* Classes, including:
    * What they are
    * Creating our own
    * Why they are important
* Polymorphism
* Objects (and instances!)

### Classes

A class is actually a type of data structure that binds variables and functions into a single, reusable unit. When we create an instance of a class, or one unit of the class, we are creating a class object.

Variables that belong to a class/class instance/object are called object properties or attributes, and functions that belong to a class/class instance/object are called object methods.

We have already seen some classes. For example, our existing data structures are built-in classes.

`list()` creates an object/instance of the `list` class. When we call `list()` and provide arguments, we are actually supplying the object parameters. When we use functions that begin with a `.`, we are actually calling a method of the object.

Below, `x` is a `list` object. `.append()` is a method. There are no accessible attributes because lists are dynamic and don't actually store data within them... but we won't concern ourselves with that.

```python
x = list((5, 6, 7))
type(x)
```

Output:
```
list
```

```python
x.append(8) # this is a method!
```

To create our own class, we start with the `class` keyword and the name of the class we are creating. Then, we create a special method, or a `constructor`, called `__init__()`.

This method is called automatically when we create an instance of our class. It always begins with `self`, which is the pointer to our object. We give any required parameters after. Then, in our constructor, we construct our class! All the attributes begin with `self` because each new object will have its own values for its attributes.

```python
class Car:
  def __init__(self, make, model, year):
    self.make = make
    self.model = model
    self.year = year

c1 = Car("Hyundai", "Venue", 2021)
c2 = Car("Chevy", "Malibu", 2010)

print(c1.make) # grabbing the value of the attribute 'make' from the c1 object
print(c1.year) # grabbing the value of the attribute 'year' from the c1 object

print(c2.make)
print(c2.model)
```

Output:
```
Hyundai
2021
Chevy
Malibu
```

Python allows other special methods, which interact with built in functions like `len()`, `print()` or operators like `+` and `**`. Take a moment to read about them here

<a href="https://www.pythonlikeyoumeanit.com/Module4_OOP/Special_Methods.html">Special Methods in Python</a>

To create a method, we just define a function within the class like usual. However, the first parameter must be `self`, and any attributes we want to access or manipulate must also start with `self`.

```python
class Car:
  def __init__(self, make, model, year):
    self.make = make
    self.model = model
    self.year = year
    self.speed = 0

  def move(self):
      print("Vroom vroom")
      self.speed = self.speed + 60

  def speedUp(self, speed_inc):
      self.speed += speed_inc

c1 = Car("Hyundai", "Venue", 2021)
c2 = Car("Chevy", "Malibu", 2010)

print(c1.make) # grabbing the value of the attribute 'make' from the c1 object
print(c1.year) # grabbing the value of the attribute 'year' from the c1 object

print(c2.make)
print(c2.model)

c1.move()
c1.speed
c1.speedUp(30)
c1.speed
```

Output:
```
Hyundai
2021
Chevy
Malibu
Vroom vroom
90
```

### Inheritance

Sometimes, we might have multiple classes that have similar attributes or methods.

```python
class Car:
  def __init__(self, make, model, year):
    self.make = make
    self.model = model
    self.year = year
    self.speed = 0

  def move(self):
      print("Vroom vroom")
      self.speed = self.speed + 60

  def speedUp(self, speed_inc):
      self.speed += speed_inc


class Airplane:
  def __init__(self, make, model, year):
    self.make = make
    self.model = model
    self.year = year
    self.speed = 0

  def move(self):
      print("Fly fly?")
      self.speed = self.speed + 60

  def speedUp(self, speed_inc):
      self.speed += speed_inc
```

In this case, we can create a parent (base) class that holds all the similar attributes and methods. Then, we can create a child (derived) class for cars, and another for airplanes. This practice is called `inheritance`.

```python
class Vehicle:
    def __init__(self, make, model, year):
        self.make = make
        self.model = model
        self.year = year
        self.speed = 0

    def move(self):
        print("Vroom vroom")
        self.speed = self.speed + 60

    def speedUp(self, speed_inc):
        self.speed += speed_inc

class Car(Vehicle):
    def __init__(self, make, model, year, units = "mph"):
        super().__init__(make, model, year)
        self.units = units

class Airplane(Vehicle):
    def __init__(self, make, model, year):
        super().__init__(make, model, year)
    def move(self):
        print("Fly fly?")
        self.speed = self.speed + 60


boeing = Airplane("Boeing", "Boeing", 2021)
boeing.move()

hyundai = Car("Hyundai", "Venue", 2021)
hyundai.move()
hyundai.speedUp(60)
print(hyundai.speed)
print(hyundai.units)
```

Output:
```
Fly fly?
Vroom vroom
120
mph
```

Notice what I did. First, the Car class has one more attribute `units` than the parent class. The parent class doesn't have this attribute, so the Airplane class does not either.

Second, we overwrote the `.move()` method of the Airplane class. Now it prints something different - but our Vehicle and Car classes remain unchanged.

Third, I introduced a default argument to the `units` parameter like I could with a regular function.

Fourth, I showed a default attribute that can't be made when initialization occurs: `speed`. It is automatically set to 0, but the function `speedUp` can increase it!

Play with the code above and see what else you can do!

### Polymorphism

Being able to use the same method with different classes is called polymorphism. We have that with the `move()` method in our example above, because `move()` can be used on both Car and Airplane objects!

Another example is the Python function `len()`. For strings, it returns the number of characters. For lists, it returns the number of items. For dictionaries, it returns the number of key/value pairs.

```python
print(len("hello!"))
print(len([1, 2, 3, 4, 5]))
print(len({'hi': 0, 'hello': 'world'}))
```

Output:
```
6
5
2
```

## Assignment Section

**Question 1.** For this question, you will create a class called `Calculator()`

It will have 6 functions:
* `__init__()` : It should define an instance variable that holds the current value of the calculator. You can name it whatever you like. The calculator should be able to be initialized with any number, but should default to 0. (FOr example, `Calculator(6)` means the starting value of the calculator is 6, while `Calculator()` means it is 0.
* `add()` : It should have a value as a parameter, and it should add that value to the current value. It should reutrn the current value.
* `subtract()` : It should have a value as a parameter, and it should subtract that value from the current value. It should return the current value.
* `multiply()` : It should have a value as a parameter, and it should multiply that value with the current value. It should return the current value.
* `divide()` : It should have a value as a parameter, and it should divide the current value by it. It should return the current value.
* `clear()` : It should reset the current value of the calculator to 0, then return the current value.

```python
class Calculator:
    def __init__(self,init_val =0):
        self.init_val = init_val
    
    def add(self, val =5):
        self.init_val += val
        return self.init_val

    def subtract(self, val =2):
        self.init_val -= val
        return self.init_val

    def multiply(self, val =10):
        self.init_val *= val
        return self.init_val

    def divide(self, val =2):
        if val == 0:
            raise ValueError("Cannot divide by zero")
        self.init_val /= val
        return self.init_val
    
    def clear(self):
        self.init_val = 0
        return self.init_val
        
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2.** For this question, you will create a class called `Vector()` with several operator overloading methods and special methods. These are methods that define custom behaviors for operators such as `==`, `+`, `abs()`, etc.

It will have 8 functions:
* `__init__()` : It should define two instance variables: `x_comp` and `y_comp`. They are given upon instantiation - there should be no default values. For example: `Vector(4, 5)` means `self.x_comp == 4` and `self.y_comp == 5`.
* `__eq__()` : It should take another Vector as a parameter (with no default). It should return `True` if the `x_comp` and `y_comp` variables match between the two instances.
* `__abs__()` : It does not take additional parameters. It should return the magnitude, or Euclidean norm, of the Vector.
* `__str__()` : It does not take additional parameters. It should return a string of the Vector. The string should be the `x_comp` variable, a space, then the `y_comp` variable, exactly.
* `__mul__()` : It should take another Vector as a parameter (with no default). It should return a new Vector. The new Vector object's `x_comp` should be the multiplication of the current Vector and the one given as an argument. The same is true for `y_comp`. For example, Vector(3, 4) * Vector(2, 3) == Vector(6, 12)
* `__add__()` : It should take another Vector as a parameter (with no default). It should return a new Vector. The new Vector object's `x_comp` should be the addition of the current Vector and the one given as an argument. The same is true for `y_comp`. For example, Vector(3, 4) + Vector(2, 3) == Vector(5, 7)
* `__sub__()` : It should take another Vector as a parameter (with no default). It should return a new Vector. The new Vector object's `x_comp` should be the subtraction of Vector given as an argument from the current Vector. The same is true for `y_comp`. For example, Vector(3, 4) - Vector(2, 3) == Vector(1, 1)
* `dot()` : It should take another Vector as a parameter (with no default). It should return a scalar. The scalar is the dot-product between the current Vector and the Vector given as an argument. (Hint: Multiply the matching variables, then add together, return the result.)

```python
class Vector:
    def __init__(self, x_comp,y_comp):
        self.x_comp = x_comp
        self.y_comp = y_comp

    def __eq__(self, new_vec):
        return self.x_comp == new_vec.x_comp and self.y_comp == new_vec.y_comp

    def __abs__(self):
        return (self.x_comp**2 + self.y_comp**2)**0.5
    
    def __str__(self):
        return (str(self.x_comp) + " " + str(self.y_comp))

    def __mul__(self, new_vec):
        new_x = self.x_comp * new_vec.x_comp
        new_y = self.y_comp * new_vec.y_comp
        return Vector(new_x, new_y)
    
    def __add__(self, new_vec):
        new_x = self.x_comp + new_vec.x_comp
        new_y = self.y_comp + new_vec.y_comp
        return Vector(new_x, new_y)
    
    def __sub__(self, new_vec):
        new_x = self.x_comp - new_vec.x_comp
        new_y = self.y_comp - new_vec.y_comp
        return Vector(new_x, new_y)
    
    def dot(self, new_vec):
        xdot = self.x_comp * new_vec.x_comp
        ydot = self.y_comp * new_vec.y_comp
        return xdot + ydot
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!
```

Tests for q1 (points: 5)
```
>>> calc = Calculator()
>>> assert calc.add(4) == 4
```
```
>>> calc = Calculator(2)
>>> assert calc.add(4) == 6
```
```
>>> calc = Calculator(4)
>>> assert calc.subtract(1) == 3
```
```
>>> calc = Calculator(3)
>>> assert calc.multiply(2) == 6
```
```
>>> calc = Calculator(6)
>>> assert calc.divide(3) == 2
```
```
>>> calc = Calculator(9385)
>>> assert calc.clear() == 0
```


Tests for q2 (points: 5)
```
>>> vect1 = Vector(4, 5)
>>> vect2 = Vector(6, 7)
>>> assert vect1.x_comp == 4
>>> assert vect1.y_comp == 5
```
```
>>> vect1 = Vector(4, 5)
>>> vect2 = Vector(6, 7)
>>> assert vect2.x_comp == 6
>>> assert vect2.y_comp == 7
```
```
>>> vect1 = Vector(4, 5)
>>> assert abs(vect1) == 41 ** 0.5
```
```
>>> vect2 = Vector(6, 7)
>>> assert abs(vect2) == 85 ** 0.5
```
```
>>> vect1 = Vector(4, 5)
>>> vect2 = Vector(6, 7)
>>> assert vect1 + vect2 == Vector(10, 12)
```
```
>>> vect1 = Vector(4, 5)
>>> vect2 = Vector(6, 7)
>>> assert vect1 - vect2 == Vector(-2, -2)
```
```
>>> vect1 = Vector(4, 5)
>>> vect2 = Vector(6, 7)
>>> assert vect1 * vect2 == Vector(24, 35)
```
```
>>> vect1 = Vector(4, 5)
>>> vect2 = Vector(6, 7)
>>> assert vect1.dot(vect2) == 59
```
```
>>> vect1 = Vector(4, 5)
>>> assert str(vect1) == '4 5'
```
